<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

# Prise en main de Jupyter Note book 

<img src="1SSPCH.png" align ="right" width = 150 height = 150>Jupyter notebook est une application web qui juxtapose sur un seul document :

* des cellules de texte (HTML, Markdown et LaTeX) pouvant recevoir des consignes, des commentaires, des questions et leurs réponses… ;
* des cellules de code pour des portions du programme à exécuter et l’affichage de son résultat sur ce même document.

L’interactivité de jupyter notebook qui permet aux élèves d’expérimenter et de tester leur code au fur à mesure de l’avancée dans leur document de travail, devrait faciliter l’apprentissage de la programmation en python et bien plus si affinité…

La cellule ci-dessous est une cellule de code, pour l'éxécuter il suffit de la sélectionner et de cliquer sur <code>Exécuter</code> ci dessus.

```python
print("Hello world!")
```

Output:
```
Hello world!
```

Le script ci-dessus est simpliste mais on peut évidemment coder des choses bien plus complexe, en fait seule votre compétence en code python sera votre limite.
Voyons ci-dessous un exemple plus complexe, qui permet d'afficher la population de Madagascar par district. 
Vous savez maintenant où se trouve le bouton <code>Exécuter</code>

```python
# fichier obtenu sur http://www.resiliencemada.gov.mg/layers/geonode:population_mada/metadata_detail
# données 2015
import folium
import webbrowser
import csv

# importation au format csv
csvfile = open('population_mada.csv')
population_csv = csv.reader(csvfile, delimiter=',',)

# Création de la carte centrée sur 45°N 1°W
carte_stations = folium.Map(location=[-18.9,47], zoom_start=6)

i = 0
for ligne in population_csv :
    if i!=0 :
        longitude = float(ligne[5])
        latitude = float(ligne[6])
        population = int(ligne[7])
        nom = str(ligne[2])
        nom = nom + " : " + str(population) + " hab"
        if population > 200000 : 
            folium.Marker([latitude,longitude], popup = nom, icon = folium.Icon(color='red')).add_to(carte_stations)
        elif population > 100000 : 
            folium.Marker([latitude,longitude], popup = nom, icon = folium.Icon(color='blue')).add_to(carte_stations)
        elif population > 50000 : 
            folium.Marker([latitude,longitude], popup = nom, icon = folium.Icon(color='green')).add_to(carte_stations)
        elif population > 30000 : 
            folium.Marker([latitude,longitude], popup = nom, icon = folium.Icon(color='orange')).add_to(carte_stations)
    else :
        i+=1

carte_stations
```

Output:
```
<folium.folium.Map at 0x7f2b267ab6a0>
```

  Vous pouvez Interargir avec les cellules, essayons de voir cela. Considérons le code ci-dessous. Il défini une fonction <code>calcul(n)</code>qui renvoie une valeur telle que $calcul(n) = 2 n + 4$, on peut le vérifier en demandant d'afficher $calcul(4)=2 \cdot 4 + 4 =12$

```python
def calcul(n) :
    return (2*n+4)
    
calcul(4)
```

Output:
```
12
```

Vous pouvez modifier la cellule ci-dessous, il suffit de cliquer dessus puis de modifier le code qu'elle contient.

Essayez de modifier la fonction <code>calcul(n)</code> de telle façon qu'elle puisse renvoyer un valeur telle que $calcul(n)=n^2-3$.

Entrez votre nouveau code dans la cellule ci-dessous, en remplaçant 
<code># YOUR CODE HERE</code>
<code>raise NotImplementedError()</code> par votre propre code

Si votre code a été correctement modifié, la cellule ci-dessous doit s'éxécuter sans renvoyer d'erreur. c'est une cellule de test qui permet de valider la modification de votre code conformément aux consignes.

<!-- grade_id: cell-f5faaf2f4c707e45, points: 0, grade: false, solution: true -->
```python
def calcul(n) :
    return ((n**2)-3)
calcul(4)
```

Output:
```
13
```

Si votre code a été correctement modifié, la cellule ci-dessous doit s'éxécuter sans renvoyer d'erreur. c'est une cellule de test qui permet de valider la modification de votre code conformément aux consignes.

<!-- grade_id: cell-5c3c06e19e960797, points: 2, grade: true, solution: false -->
```python
assert calcul(0) == -3
assert calcul(2) == 1
assert calcul(-1) == -2
```

## Maintenant à vous !!
Vous pouvez maintenant fermer ce NoteBook, et passer au suivant qui vous proposera quelques exercices *_2.Premiers exercices Notebook._*

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>

```python

```

```python

```

```python

```