<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

# Premier exercice en python
Vous allez avec ce NoteBook prendre connaissance sur un exercice très simple de la procédure à suivre pour rendre les exercices demandés.
Cette procédure sera toujours décrite dans la cellule d'en-tête, ci-dessus mais voyons la ici en détail.
Il faut toujours être attentif à l'ordre d'éxécution des cellules.

## 1. Texte de l'exercice, et problème à résoudre.
Ecrire une fonction <code>racine(n)</code> qui renvoie la racine carré d'un nombre positif ou nul, et qui renvoie <code>None</code> sinon. On doit donc avoir $racine(n) = \sqrt n$
Vous allez donc devoir effacer :

<code># YOUR CODE HERE
raise NotImplementedError()</code>

et le remplacer par votre propre code.

<!-- grade_id: cell-99519086d872aa9d, points: 0, grade: false, solution: true -->
```python
from math import sqrt
def racine(n) :
    if n>=0 :
        return sqrt(n)
    else : 
        print ("La racine carré d'un nombre négatif n'existe pas.")
racine(-16)
```

Output:
```
La racine carré d'un nombre négatif n'existe pas.
```

## 2. Tests Unitaires et Jeu d'essai
Pour vérifier que votre fonction est correctement écrite, nous allons lui faire passer maintenant un jeu de tests. Certains de ces tests sont visibles, mais attention certains sont cachés et visible uniquement par l'enseignant.
Généralement ce sont ces tests qui vont vous rapporter des points si l'exercice est noté.
**Pour vérifier que tout va bien, la cellule ci-dessous doit s'éxécuter sans générer d'erreurs.**

<!-- grade_id: cell-24378c4e9822da56, points: 2, grade: true, solution: false -->
```python
assert racine(4) == 2
assert racine(-2) == None
```

Output:
```
La racine carré d'un nombre négatif n'existe pas.
```

## 3. Soumettre votre exercice
Une fois votre exercice terminé, dans l'ideal votre fonction a réussi à passer tous les tests, en tout cas au moins les test visibles.
Il ne vous reste plus qu'à soumettre votre travail :
1. Pensez à enregistrer votre NoteBook et revenez sur la page de gestions de fichier (**icône Jupyter**)
2. Cliquez sur l'ongler <code>Assignements</code>.
3. Cliquer sur le bouton <code>Submit</code>.
4. Votre fichier est envoyé pour évaluation à votre professeur.

Remarque : Vous pouvez envoyer plusieurs <code>Submit</code> mais attention seul le dernier sera évalué, à condition qu'il soit envoyé dans les délais imposés par l'exercice.

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>